<problem>
<script type="loncapa/python">
<![CDATA[
# This Python code defines and evaluates the problem. You usually don't need to touch the
# first portion of the code and can jump to "Problem definition" near the end.

import decimal
import json


class ResponseCell(object):
    """Abstract base class for response cells."""

    TYPE = 0

    def get_state(self):
        """Return a dictionary describing the instance for sending to the JS code."""
        d = vars(self).copy()
        d.update(type=self.TYPE)
        return d


class Numeric(ResponseCell):
    """A numeric response cell."""

    TYPE = 1

    def __init__(self, answer, tolerance=None,
                 min_significant_digits=None, max_significant_digits=None):
        """Set the correct answer and the allowed relative tolerance in percent."""
        self.answer = answer
        if tolerance is None:
            tolerance = default_tolerance
        self.abs_tolerance = abs(answer) * tolerance / 100.0
        self.min_significant_digits = min_significant_digits;
        self.max_significant_digits = max_significant_digits;

    def check_response(self, student_response):
        """Return a Boolean value indicating whether the student response is correct."""
        try:
            r = float(student_response)
        except ValueError:
            return False
        if self.min_significant_digits or self.max_significant_digits:
            d = len(decimal.Decimal(student_response).as_tuple().digits)
            if self.min_significant_digits and d < self.min_significant_digits:
                return False
            if self.max_significant_digits and d > self.max_significant_digits:
                return False
        return abs(r - self.answer) <= self.abs_tolerance


class String(ResponseCell):
    """A string response cell."""

    TYPE = 2

    def __init__(self, answer):
        """Set the correct answer."""
        self.answer = answer

    def check_response(self, student_response):
        """Return a Boolean value indicating whether the student response is correct."""
        return student_response == self.answer


def init(table, help_text, column_widths, row_height):
    """Initialisation code for values exported to the XML code.

    The initial_state is a JSON string exported to the JavaScript code and used to construct
    the table.  The height is the height of the iframe in pixels.
    """
    data = [
        [cell.get_state() if isinstance(cell, ResponseCell) else str(cell) for cell in row]
        for row in table
    ]
    initial_state = dict(
        data=data,
        help_text=help_text,
        column_widths=column_widths,
        row_height=row_height,
    )
    # Curiously, the JSON string must be doubly HTML-escaped.
    initial_state = json.dumps(initial_state).replace('&', '&amp;amp;').replace('"', '&amp;quot;')
    # All heights and widths include a bit of safety margin to avoid scrollbars and rendering
    # artifacts on some browsers.
    if column_widths is not None:
        width = sum(column_widths) + 2
    else:
        width = 800 + 2
    height = (row_height + 1) * len(table)
    if help_text:
        height += (len(help_text) // 100 + 2) * 26
    return initial_state, width, height


def check_consistency(state):
    """Check whether the state of the table is consistent with the initial state.

    If not, the problem was edited in the mean time, and we ask the student to click
    the reset button (as we can't initiate resetting the problem from code).
    """
    if len(table) != len(state['data']):
        return False
    for row, row_state in zip(table, state['data']):
        if len(row) != len(row_state):
            return False
        for cell, cell_state in zip(row, row_state):
            if isinstance(cell, ResponseCell):
                if not isinstance(cell_state, dict):
                    return False
                orig_state = cell.get_state()
                if not all(orig_state[k] == cell_state[k] for k in orig_state):
                    return False
            else:
                if not str(cell) == cell_state:
                    return False
    return True


def check_table(unused_expect, ans):
    """Determine the correctness of all student responses.

    The correctness is independently determined by the JavaScript code on the client side.
    """
    ans = json.loads(ans)
    state = json.loads(ans['state'])

    if not check_consistency(state):
        return {
            'ok': False,
            'msg': 'The problem was edited by the course staff since the last time you '
            'worked on it.  Please click the "Reset" button to get the latest version.'
        }

    correct_answers = total_questions = 0
    for row, row_state in zip(table, state['data']):
        for cell, cell_state in zip(row, row_state):
            if isinstance(cell, ResponseCell):
                total_questions += 1
                correct_answers += cell.check_response(cell_state['value'])

    if correct_answers == total_questions:
        return {'ok': True, 'msg': 'Great job!'}
    else:
        return {
            'ok': False,
            'msg': 'You have {} cells out of {} correct.'.format(correct_answers, total_questions)
        }

## Problem definition ##
## Edit the table and settings below. ##

# The default tolerance in percent, used for all numeric fields that don't specify an ecplicit
# tolerance.
default_tolerance = 1.0

# The table itself.  The first row contains the column headers.  Any cell in the following rows can
# be either a string or number, which will simply be shown to the student, or one of the special
# cell types Numeric() or String(), indicating a numeric or string input field that must be filled
# in by the student.
#
# Syntax of the special cell types:
#
#     Numeric(answer=<correct_answer>, tolerance=<tolerance in percent>,
#             min_significant_digits=<number>, max_significant_digits=<number>)
#         The tolerance is optional, and will default to the default tolerance specified above.
#         The restrictions for the number of significant digits are optional as well.  Significant
#         digits are counted started from the first non-zero digit specified by the student, and
#         include trailing zeros.
#
#     String(answer='<correct answer>')
table = [
    ['Carts', 'v_initial', 'v_final', 'Delta v'],
    ['green', 2.0, -1.2, Numeric(answer=-3.2, tolerance=2.0)],
    ['red', -2.0, 1.2, Numeric(answer=3.2)],
    ['yellow', -1.4, 0.0, Numeric(answer=1.4)],
    ['purple', 3.4, 2.1, Numeric(answer=-1.3)],
]

# Set the text that gets displayed when clicking the "+help" button.  Use "None" to disable the
# help feature.
#help_text = None
help_text = "Fill in the yellow fields with the right values."

# Set the width of the columns in pixels.  The total width of the table shouldn't be more than 800.
# You can use "column_width = None" for equal-width columns with a total width of 800 pixels.
column_widths = None
#column_widths = [150, 150, 150, 150]

# The height of a table row in pixels.  Rows may grow higher than the specified value if the text
# in some cells in the row is long enough to get wrapped in more than one line.  If this happens,
# you might need to replace "$height" close to the end of this file by a big enough value to
# prevent a vertical scrollbar for the iframe from appearing.
row_height = 36

# Initialisation -- don't remove!
initial_state, width, height = init(table, help_text, column_widths, row_height)

]]>
</script>
<p>
  This is an ActiveTable problem, allowing you to place input fields in the cell
  of a table.  You can edit the table and declare the cells that are input cells
  by clicking EDIT above and following the instructions in the Python code.
</p>
<p>
  It is recommended to set "Randomization" to "On Reset" and "Show Reset Button"
  to "True" in the settings, even when the problem does not involve any
  randomization.  This will ensure that students can correctly reload a changed
  version of a problem they've already worked on after a member of staff edited
  the problem.  You can change the problem settings by clicking EDIT above, then
  SETTINGS in the top right corner of the edit window.
</p>
<customresponse cfn="check_table">
  <jsinput gradefn="ActiveTable.grade"
    get_statefn="ActiveTable.getState"
    set_statefn="ActiveTable.setState"
    initial_state="$initial_state"
    width="$width"
    height="$height"
    html_file="/static/activetable.html"
    sop="true"/>
</customresponse>
<solution>
  <p>
    Delta_v is computed as v_final – v_initial.
  </p>
</solution>
</problem>
