<problem>
<script type="loncapa/python">
<![CDATA[
# This Python code defines and evaluates the problem. You usually don't need to touch the
# first portion of the code and can jump to "Problem definition" near the end.
import json
class ResponseCell(object):
    """Abstract base class for response cells."""
    TYPE = 0
    def get_state(self):
        """Return a dictionary describing the instance for sending to the JS code."""
        d = vars(self).copy()
        d.update(type=self.TYPE)
        return d
class Numeric(ResponseCell):
    """A numeric response cell."""
    TYPE = 1
    def __init__(self, answer, tolerance=None):
        """Set the correct answer and the allowed relative tolerance in percent."""
        self.answer = answer
        if tolerance is None:
            tolerance = default_tolerance
        self.abs_tolerance = abs(answer) * tolerance / 100.0
    def check_response(self, student_response):
        """Return a Boolean value indicating whether the student response is correct."""
        try:
            r = float(student_response)
        except ValueError:
            return False
        return abs(r - self.answer) <= self.abs_tolerance
class String(ResponseCell):
    """A string response cell."""
    TYPE = 2
    def __init__(self, answer):
        """Set the correct answer."""
        self.answer = answer
    def check_response(self, student_response):
        """Return a Boolean value indicating whether the student response is correct."""
        return student_response == self.answer
def init(table, help_text, column_widths, row_height):
    """Initialisation code for values exported to the XML code.
    The initial_state is a JSON string exported to the JavaScript code and used to construct
    the table.  The height is the height of the iframe in pixels.
    """
    data = [
        [cell.get_state() if isinstance(cell, ResponseCell) else str(cell) for cell in row]
        for row in table
    ]
    initial_state = dict(
        data=data,
        help_text=help_text,
        column_widths=column_widths,
        row_height=row_height,
    )
    # Curiously, the JSON string must be doubly HTML-escaped.
    initial_state = json.dumps(initial_state).replace('&', '&amp;amp;').replace('"', '&amp;quot;')
    # All heights and widths include a bit of safety margin to avoid scrollbars and rendering
    # artifacts on some browsers.
    if column_widths is not None:
        width = sum(column_widths) + 2
    else:
        width = 800 + 2
    height = (row_height + 1) * len(table)
    if help_text:
        height += (len(help_text) // 100 + 2) * 26
    return initial_state, width, height
def check_consistency(state):
    """Check whether the state of the table is consistent with the initial state.
    If not, the problem was edited in the mean time, and we ask the student to click
    the reset button (as we can't initiate resetting the problem from code).
    """
    if len(table) != len(state['data']):
        return False
    for row, row_state in zip(table, state['data']):
        if len(row) != len(row_state):
            return False
        for cell, cell_state in zip(row, row_state):
            if isinstance(cell, ResponseCell):
                if not isinstance(cell_state, dict):
                    return False
                orig_state = cell.get_state()
                if not all(orig_state[k] == cell_state[k] for k in orig_state):
                    return False
            else:
                if not str(cell) == cell_state:
                    return False
    return True
def check_table(unused_expect, ans):
    """Determine the correctness of all student responses.
    The correctness is independently determined by the JavaScript code on the client side.
    """
    ans = json.loads(ans)
    state = json.loads(ans['state'])
    if not check_consistency(state):
        return {
            'ok': False,
            'msg': 'The problem was edited by the course staff since the last time you '
            'worked on it.  Please click the "Reset" button to get the latest version.'
        }
    correct_answers = total_questions = 0
    for row, row_state in zip(table, state['data']):
        for cell, cell_state in zip(row, row_state):
            if isinstance(cell, ResponseCell):
                total_questions += 1
                correct_answers += cell.check_response(cell_state['value'])
    if correct_answers == total_questions:
        return {'ok': True, 'msg': 'Great job!'}
    else:
        return {
            'ok': False,
            'msg': 'You have {} cells out of {} correct.'.format(correct_answers, total_questions)
        }
## Problem definition ##
## Edit the table and settings below. ##
# The default tolerance in percent, used for all numeric fields that don't specify an ecplicit
# tolerance.
default_tolerance = 1.0
# The table itself.  The first row contains the column headers.  Any cell in the following rows can
# be either a string or number, which will simply be shown to the student, or one of the special
# cell types Numeric() or String(), indicating a numeric or string input field that must be filled
# in by the student.
# Syntax of the special cell types:
#     Numeric(answer=<correct_answer>, tolerance=<tolerance in percent>)
#     String(answer='<correct answer>')
table = [
    ['Reserve Ratio','Multiplier','New Demand Deposits','Change in Circulating Currency','Maximum Change in the Money Supply'],
    ['10%', Numeric(answer=10),Numeric(answer=30000),Numeric(answer=-3000),Numeric(answer=27000)],
    ['20', Numeric(answer=5), Numeric(answer=15000), Numeric(answer=-3000), Numeric(answer=12000)],
    ['25', Numeric(answer=4), Numeric(answer=12000), Numeric(answer=-3000), Numeric(answer=9000)],
]

# Set the text that gets displayed when clicking the "+help" button.  Use "None" to disable the
# help feature.
#help_text = None
help_text = "Fill in the yellow fields with the right values."
# Set the width of the columns in pixels.  The total width of the table shouldn't be more than 800.
# You can use "column_width = None" for equal-width columns with a total width of 800 pixels.
column_widths = None
#column_widths = [150, 150, 150, 150]
# The height of a table row in pixels.  Rows may grow higher than the specified value if the text
# in some cells in the row is long enough to get wrapped in more than more lines.  If this happens,
# you might need to replace "$height" close to the end of this file by a big enough value to
# prevent a vertical scrollbar for the iframe from appearing.
row_height = 36
# Initialisation -- don't remove!
initial_state, width, height = init(table, help_text, column_widths, row_height)
]]>
</script>
<img src="/static/EdX_Banners-12.png" width="901" height="112"/>
  <p>For each of the following scenarios, determine the multiplier, new demand deposits, the change in circulating currency, and the maximum possible change in the money supply. Consider that circulating currency of $3,000 is deposited at the First National Bank of Econoland. Assume all excess reserves are loaned out and there are no currency drains.</p>
  <p>Complete the chart below by inputting numeric responses in the input fields (do not use any commas). You may want to work through this problem on a blank piece of paper first.</p>

<p><i>Fill in the active table with your responses. Note, all input values should use decimal notation where appropriate.</i></p>
<customresponse cfn="check_table">
  <jsinput gradefn="ActiveTable.grade"
    get_statefn="ActiveTable.getState"
    set_statefn="ActiveTable.setState"
    initial_state="$initial_state"
    width="$width"
    height="$height"
    html_file="/static/activetable.html"
    sop="true"/>
</customresponse>
  <solution>
    <!--<img src="/static/moneycreation4.1.1.png" />-->
  </solution>
</problem>